ELEC-E5550 - Statistical Natural Language Processing
# SET 6: Subword segmentation

## Released: 09.03.2021 at 21:30
## Deadline: 22.03.2020 at midnight

# Overview

We've already talked about the problem of segmenting text into appropriate units (tokenization). Back then, it was words that were considered as those units, but there are other units that you should probably explore as well: characters and **subwords**. In this assignment, we're going to focus on **tokenization into subwords**.

The motivation to segmenting words further into smaller elements comes from morphology, where such elements are called **morphemes**. A **morpheme** is defined as the smallest meaning-bearing unit of a language. For example, the word *unpredictable* contains three morphemes: *un*, *predict* and *able*. As you can see, many morphemes are not unique to one word, they are elements that are regularly seen in other words too. For example, *un* in *unhappy*, *predict* in *predictive*, and *able* in *comfortable*. Thus, just as sentences are constructed from words, words are constructed from morphemes. 

Segmenting into morphemes (especially in languages with rich morphology) helps to avoid the problem of out-of-vocabulary (OOV) words in text corpora. For example, if our training corpus contains *cool*, *cool-est* and *dumb-er* when the new word *cooler* comes in, it can be analyzed into *cool-er*.

# Table of contents

* [Task 1: BPE](#task_1)
    * [Step 1.1](#subtask_1_1)
        * [Step 1.1.1: Collecting word counts](#subtask_1_1_1)
        * [Step 1.1.2: Convering words to characters](#subtask_1_1_2)
    * [Step 1.2: Collecting symbol pairs frequencies](#subtask_1_2)
    * [Step 1.3: Merging the most frequent pair](#subtask_1_3)
    * [Step 1.4: Combining steps 1-3 together](#subtask_1_4)
    * [Step 1.5-7: Segmenting a corpus](#subtask_1_5)
* [Task 2: ANALYZE SEGMENTATIONS](#task_2)
    * [Step 2.1: Count word OOV](#subtask_2_1)
    * [Step 2.2: Count subword OOV](#subtask_2_2)
    * [Step 2.3: Does the segmentation make sense?](#subtask_2_3)
    * [Step 2.4: Your thoughts](#subtask_2_4)
* [Checklist](#checklist)

## TASK 1 <a class="anchor" id="task_1"></a>
## BPE
One approach to tokenization into subwords is based on the **byte pair encoding** (**BPE**) algorithm for text compression. It iteratively merges frequent pairs of characters forming new subwords. The intuition here is that morphemes are frequently repeated substrings, so this method should merge symbols into them instead of into some random meaningless character sequences. The algorithm is applied only inside words (there is no merges across word boundaries). 

**BPE** algorithm begins with its vocabulary being a set of characters seen in the training corpus. Each word in the corpus is represented as a sequence of characters plus a special end-of-word symbol '_'. At each iteration step $k_i$, the algorithm counts the number of symbol pairs, finds the most frequent pair ['A','B'] and replaces it with the new merged symbol ‘AB’. The algorithm stops when it's done $k$ merges ($k$ is a parameter of the algorithm). The algorithm begins with the set of symbols equal to the set of characters. The resulting symbol set should have the original set of characters plus $k$ new symbols. 

To learn segmentations with **BPE**, you should take the following steps:
* STEP 1: tokenize a training corpus into words and collect frequency statistics of word tokens in the training corpus. Additionally, represent each word as a list of characters plus a special end-of-word symbol '_'.
* STEP 2: count the frequencies of symbol pairs.
* STEP 3: replace every occurrence of the most frequent pair ['A', 'B'] with the new merged symbol 'AB'.
* STEP 4: repeat STEPs 2-3 $k-1$ times more.

To segment a test corpus with the learned segmentation, you should:
* STEP 5: tokenize a test corpus into words (with the same tokenization algorithm as in training).
* STEP 6: represent each word as a list of characters plus a special end-of-word symbol '_'.
* STEP 7: for every word apply each merge operation in the order they were learned, and return the tokenized text

## 1.1 <a class="anchor" id="subtask_1_1"></a>
### STEP 1
### 1.1.1  <a class="anchor" id="subtask_1_1_1"></a>
### Collecting word counts (1 Point)

Write a function that reads a text as one string from a file, tokenizes it into words by whitespaces and collects the word frequencies. It should return a dictionary of words and their raw counts in a corpus.

<!-- grade_id: cell-1b45cebfbb7868fc, points: 0, grade: false, solution: true -->
```python
from collections import Counter
def collect_word_counts(file_name):
    """
    Takes in a path to a text file, reads the file, splits it into words by whitespaces, 
    and then counts the words' frequencies.
    
    Parameters
    ---------
    file_name : string
            a path to a training corpus as a string
    
    Returns
    -------
    word_counts : dcitionary
            a dictionary of word counts
    """
    
    # YOUR CODE HERE
    with open(file_name) as f:
        text=f.read()
        words=text.split()
        f.close()
    word_counts=dict(Counter(words))
    return word_counts
```

<!-- grade_id: cell-ac8d751900ffa327, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal

# CHECKING THE GENERAL PROPERTIES OF THE OUTPUT
dummy_corpus_path = "/coursedata/06-subwords/dummy_corpus.txt"

# check that the output of the function is a list
assert_equal(type(collect_word_counts(dummy_corpus_path)), dict)

# CHECKING THAT THE FUNCTION IS WORKING AS IT SHOULD
assert_equal(collect_word_counts(dummy_corpus_path)['low'], 5)
assert_equal(collect_word_counts(dummy_corpus_path)['lowest'], 2)
assert_equal(collect_word_counts(dummy_corpus_path)['new'], 2)
assert_equal(collect_word_counts(dummy_corpus_path)['newer'], 6)
assert_equal(collect_word_counts(dummy_corpus_path)['WIDer'], 3)


# A SANITY CHECK FOR THE NOTEBOOK DATA
gum_train_path = "/coursedata/06-subwords/gum_train.txt"
# check that the vocabulary length is right
assert_equal(len(collect_word_counts(gum_train_path)), 8677)
# check that the word count is right
assert_equal(collect_word_counts(gum_train_path)['we'], 112)


```

### 1.1.2  <a class="anchor" id="subtask_1_1_1"></a>
### Converting words to characters (1 Point)
Now, represent each word in your frequency dictionary as a tuple of characters plus a special end-of-word marker '_'.

<!-- grade_id: cell-e358614034bbeaad, points: 0, grade: false, solution: true -->
```python
def convert_to_chars(vocab):
    """
    Takes in a frequeny dictionary of words in the training corpus
    and converts the key words to a tuple of characters plus a special end-of-word symbol '_'.
    
    Parameters
    ----------
    vocab : dictionary
        a frequency dictionary of words
    
    Returns
    -------
    separated_vocab : dictionary 
        a frequency dictionary of words represented as a tuple of characters 
        plus a special end-of-word symbol '_'
        {('l','o','w','_') : 3}
    """
    
    # YOUR CODE HERE
    separated_vocab={}
    for key, value in vocab.items():
        split_key=list(key)
        split_key.append('_')
        separated_vocab[tuple(split_key)]=value
    return separated_vocab
```

<!-- grade_id: cell-8efbbf91eee65873, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal

# CHECKING THE GENERAL PROPERTIES OF THE OUTPUT
dummy_freq_vocab = {'low': 5, 'lowest': 2, 'newer': 6, 'wider': 3, 'New': 2}

# check that the output of the function is a dict
assert_equal(type(convert_to_chars(dummy_freq_vocab)), dict)
# check that the keys are tuples
assert_equal(type(list(convert_to_chars(dummy_freq_vocab).keys())[0]), tuple)
#check that there are no new elements
assert_equal(len(convert_to_chars(dummy_freq_vocab)), 5)

# CHECKING THAT THE FUNCTION IS WORKING AS IT SHOULD

assert_equal(convert_to_chars(dummy_freq_vocab)[('l', 'o', 'w', '_')], 5)
assert_equal(convert_to_chars(dummy_freq_vocab)[('l', 'o', 'w', 'e', 's', 't', '_')], 2)
assert_equal(convert_to_chars(dummy_freq_vocab)[('n', 'e', 'w', 'e', 'r', '_')], 6)
assert_equal(convert_to_chars(dummy_freq_vocab)[('w', 'i', 'd', 'e', 'r', '_')], 3)
assert_equal(convert_to_chars(dummy_freq_vocab)[('N', 'e', 'w', '_')], 2)


```

## 1.2 <a class="anchor" id="subtask_1_2"></a>
### STEP 2
### Collecting symbol pairs frequencies (1 Point)
Write a function that takes in a frequency dictionary where keys are words represented as tuples of symbols, and outputs the most frequent pair of symbols in the corpus. In the case, when there are several pairs with the same frequency, return the pair that is earlier alphabetically. 

For example, if we only have one *l o o k _* and one *l o o p _*  in our frequency dictionary, the fuction should output *l o* as the most frequent pair (it is earlier than *o o* alphabetically).

<!-- grade_id: cell-e2a9fb6eb0db7f82, points: 0, grade: false, solution: true -->
```python
def get_the_pair_to_merge(vocab_as_symbols):
    """
    Takes in a frequency dictionary, where keys are words represented as tuples of symbols and values are their counts, 
    and outputs the most frequent pair of symbols in the corpus.
    
    Arguments
    ---------
    vocab_as_symbols : dictionary
        a frequency dictionary, where keys are words represented as tuples of symbols and values are their counts
    
    Returns
    -------
    merge_pair : tuple of strings 
        the most frequent pair of symbols, a pair to merge
    """
    
    # YOUR CODE HERE
    statistic={}
    max_value=0
    for key, value in vocab_as_symbols.items():

        for index,char in enumerate(key):
            if index<len(key)-1:
                pair=(char,key[index+1])
                if pair in statistic.keys():
                    statistic[pair]+=value
                else:
                    statistic[pair]=value

                if statistic[pair]> max_value:
                        max_value=statistic[pair]
                        merge_pair=pair
                elif statistic[pair]== max_value and pair<merge_pair:
                    max_value=statistic[pair]
                    merge_pair=pair        

    
    return merge_pair
```

<!-- grade_id: cell-74bc7d2d9ec1da10, points: 1, grade: true, solution: false -->
```python
# CHECKING THE GENERAL PROPERTIES OF THE OUTPUT
dummy_freq_vocab_as_symbols = {('l','o','o','k','_') : 1, 
                               ('l','o','o','p','_') : 1}

# check that the output of the function is a tuple
assert_equal(type(get_the_pair_to_merge(dummy_freq_vocab_as_symbols)), tuple)
# check that the output of the function is a tuple of strings
assert_equal(type(get_the_pair_to_merge(dummy_freq_vocab_as_symbols)[0]), str)


# CHECKING THAT THE FUNCTION IS WORKING AS IT SHOULD


assert_equal(get_the_pair_to_merge(dummy_freq_vocab_as_symbols), ('l', 'o'))


dummy_freq_vocab_as_symbols2 = {('l', 'o', 'w', '_'): 5,
                                ('l', 'o', 'w', 'e', 's', 't', '_'): 2,
                                ('n', 'e', 'w', 'e', 'r', '_'): 6, 
                                ('w', 'i', 'd', 'e', 'r', '_'): 3, 
                                ('n', 'e', 'w', '_'): 2}


assert_equal(get_the_pair_to_merge(dummy_freq_vocab_as_symbols2), ('e', 'r'))

# check that your functions can merge not only characters

dummy_freq_vocab_as_symbols3 = {('lo', 'o', 'k', '_'): 1 , 
                                ('lo', 'o', 'p','_'): 1}

assert_equal(get_the_pair_to_merge(dummy_freq_vocab_as_symbols3), ('lo', 'o'))

```

## 1.3 <a class="anchor" id="subtask_1_3"></a>
### STEP 3
### Merging the most frequent pair (3 Points)

Write a function that takes in a pair of symbols to merge and a frequency dictionary where words are represented as tuples of symbols, and returns a frequency dictionary, where words are still represented as tuples of symbols, but the most frequnt pair is now merged in every word.


For example, if we want to merge *l* and *o* in our *l o o k _* and *l o o p _*  frequency dictionary, the fuction should output *lo o k _* and *lo o p _*

<!-- grade_id: cell-7f9d1e7549722c1c, points: 0, grade: false, solution: true -->
```python
def merge(vocab_as_symbols, merge_pair):
    """Merges the most frequent pair of symbols in all words
    
    Parameters
    ---------
    vocab_as_symbols : dictionary
        a frequency dictionary, where keys are words represented as tuples of symbols and values are their counts
    merge_pair : tuple
        a pair of symbols to merge
        
    Returns
    -------
    new_vocab : dictionary
        a frequency dictionary, where keys are words represented as tuples of symbols, 
        with the given pair represented as a new symbol (concatenated pair)
    """
    
    # YOUR CODE HERE
    new_vocab={}
    
    for key, value in vocab_as_symbols.items():
        newkey=()
        skip_index=-1
        for index,char in enumerate(key):
            if index==skip_index:
                continue
            if index<len(key)-1 and (char, key[index+1]) == merge_pair:
                newkey+=(merge_pair[0]+merge_pair[1],)
                skip_index=index+1
            else:
                newkey+=(char,)
        new_vocab[newkey]=value
    return new_vocab
```

<!-- grade_id: cell-e71ea4c2a2aee8c2, points: 3, grade: true, solution: false -->
```python
from nose.tools import assert_equal

# CHECKING THE GENERAL PROPERTIES OF THE OUTPUT
dummy_vocab_as_symbols = {('l', 'o', 'w', '_'): 5,
                          ('l', 'o', 'w', 'e', 's', 't', '_'): 2,
                          ('n', 'e', 'w', 'e', 'r', '_'): 6, 
                          ('w', 'i', 'd', 'e', 'r', '_'): 3, 
                          ('n', 'e', 'w', '_'): 2}

dummy_pair = ('e', 'r')

# check that the output of the function is a dictionary
assert_equal(type(merge(dummy_vocab_as_symbols, dummy_pair)), dict)
# check that the keys are tuples
assert_equal(type(list(merge(dummy_vocab_as_symbols, dummy_pair).keys())[0]), tuple)
#check that there are no new elements
assert_equal(len(merge(dummy_vocab_as_symbols, dummy_pair)), 5)

# CHECKING THAT THE FUNCTION IS WORKING AS IT SHOULD
# check that the pair was merged everywhere
assert_equal(list(merge(dummy_vocab_as_symbols, dummy_pair).keys()), [('l', 'o', 'w', '_'), 
                                                                    ('l', 'o', 'w', 'e', 's', 't', '_'), 
                                                                    ('n', 'e', 'w', 'er', '_'), 
                                                                    ('w', 'i', 'd', 'er', '_'), 
                                                                    ('n', 'e', 'w', '_')])
# check that you can handle several merge pairs in a word
dummy_vocab_as_symbols2 = {('l','o','o','o','l'): 1, 
                           ('l','o','o','o','o','l'):1}

dummy_pair2 = ('o', 'o')


assert_equal(list(merge(dummy_vocab_as_symbols2, dummy_pair2).keys()), [('l', 'oo', 'o', 'l'), 
                                                                        ('l','oo','oo','l')])

```

## 1.4 <a class="anchor" id="subtask_1_4"></a>
### STEP 4
### Combining steps 1-3 together (1 Point)
Now let's combine steps 1-3 into a function that learns $k$ BPE merges.

<!-- grade_id: cell-ccfb0b9de81867b7, points: 0, grade: false, solution: true -->
```python
def learn_BPE_merges(file_name, k):
    """ Learns k BPE subwords
    
    STEP 1: Collect a word count dictionary from a file
            represent words as a tuple of their characters plus a special end-of-word symbol '_'
    Now k times
        STEP 2: Choose the most frequent pair of symbols
                add this pair as a new subword unit into a subword vocabulary
        STEP 3: Merge the symbols in all words

        
    
    Parameters
    ---------
    file_name : string
        a path to a training corpus as a string
    k : integer
        a number of merges to be learned
    
    Returns
    -------
    merges : list of strings
        a list of k subwords (symbol merges) in the order they were learned
        one merge is a tuple of two symbols in the most frequent pair at step k
    """
    
    merges = []
    # YOUR CODE HERE
    word_counts=collect_word_counts(file_name)
    vocab_as_symbols=convert_to_chars(word_counts)
    for _ in range(k):
        
        merge_pair=get_the_pair_to_merge(vocab_as_symbols)
        vocab_as_symbols=merge(vocab_as_symbols, merge_pair)
        merges.append(merge_pair)
    return merges
```

<!-- grade_id: cell-5bfcbdbe63a3fdff, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal

# CHECKING THE GENERAL PROPERTIES OF THE OUTPUT
dummy_corpus_path = "/coursedata/06-subwords/dummy_corpus.txt"


# check that the output of the function is a list
assert_equal(type(learn_BPE_merges(dummy_corpus_path, 10)), list)
# check that the output of the function is a list of tuples
assert_equal(type(learn_BPE_merges(dummy_corpus_path, 10)[0]), tuple)
# check that the output of the function is a list of tuples of strings
assert_equal(type(learn_BPE_merges(dummy_corpus_path, 10)[0][0]), str)
#check that there are exactly k merges
assert_equal(len(learn_BPE_merges(dummy_corpus_path, 10)), 10)

# CHECKING THAT THE FUNCTION IS WORKING AS IT SHOULD
# check that 8 dummy merges are correct
assert_equal(learn_BPE_merges(dummy_corpus_path, 8), [('e', 'r'),
                                                     ('er', '_'),
                                                     ('e', 'w'),
                                                     ('n', 'ew'),
                                                     ('l', 'o'),
                                                     ('lo', 'w'),
                                                     ('new', 'er_'),
                                                     ('low', '_')])

# A SANITY CHECK FOR THE NOTEBOOK DATA
gum_train_path = "/coursedata/06-subwords/gum_train.txt"
assert_equal(learn_BPE_merges(gum_train_path, 5), [('e', '_'),
                                                   ('s', '_'),
                                                   ('t', 'h'),
                                                   ('t', '_'),
                                                   ('d', '_')])

             
```

## 1.5-7 <a class="anchor" id="subtask_1_5"></a>
### STEPS 5-7
### Segmenting a corpus (5 Points)
Well, now we can apply what we've learned to tokenize any text into its subwords. Write a function that reads the test corpus and applies the merges we've learned on a training corpus. Note that you will need to adapt your previous functions a little for this. For example, we don't need to count the word frequencies since they don't play any role here anymore.

Just a reminder of the steps needed to apply a subword tokenzation to a new (or an old) text:

* STEP 5: tokenize a test corpus into words (with the same tokenization algorithm as in training).
* STEP 6: represent each word as a list of characters plus a special end-of-word symbol '_'.
* STEP 7: for every word apply each merge operation in the order they were learned, and return the tokenized text.


For the purposes of the exercise, the tokenized text should be a list of strings where strings are words with their subwords separated by whitespaces: ['I', 'lo ok', 'g oo d']. This way it will be easier for you to check how each word is tokenized. In the real application, a tokenized text will be represented just as a list of subwords.


Note: don't forget to get rid of the special end-of-word symbol after tokenization!

<!-- grade_id: cell-bc393db75acac1eb, points: 0, grade: false, solution: true -->
```python
def segment_text(file_name, merges):
    """
    Takes in a path to a text file and lerned BPE merges,
    reads the file and tokenizes it into subwords in accorance with the merges.
    
    Arguments
    ---------
    file_name : string
        a path to a text as a string
    merges : list of tuples
        a list of k merges in the order they were learned
    
    Returns
    -------
    segmented_text - list of strings
        a text segmented with BPE
        the text is a list of words
        where each word is a string with its segments separated by whitespaces:
        ['I', 'lo ok', 'g oo d']
    """
    
    # YOUR CODE HERE
    segmented_text=[]
    with open(file_name) as f:
        text=f.read()
        words=text.split()
        f.close()
    
    for word in words:
        word_list=list(word)
        word_list.append('_')
        for merge in merges:
            new_word_list=[]
            skip_index=-1
            for index ,char in enumerate(word_list):
                if skip_index==index:
                    continue
                if index < len(word_list)-1 and (char,word_list[index+1])==merge:
                    new_word_list.append(char+word_list[index+1])
                    skip_index=index+1
                else:
                    new_word_list.append(char)
            word_list=new_word_list
        if '_' in word_list:
            new_word=' '.join(word_list[:-1])
        else:
            new_word=' '.join(word_list).split('_')[0]
        segmented_text.append(new_word)
    return segmented_text
```

<!-- grade_id: cell-6d91df8425edb415, points: 5, grade: true, solution: false -->
```python
from nose.tools import assert_equal

# CHECKING THE GENERAL PROPERTIES OF THE OUTPUT
dummy_train_path = "/coursedata/06-subwords/dummy_corpus.txt"

dummy_merges = [('e', 'r'),
                ('er', '_'),
                ('e', 'w'),
                ('n', 'ew'),
                ('l', 'o'),
                ('lo', 'w'),
                ('ma', 'ma')]


# check that the output of the function is a list
assert_equal(type(segment_text(dummy_train_path, dummy_merges)), list)
# check that the output of the function is a list of strings
assert_equal(type(segment_text(dummy_train_path, dummy_merges)[0]), str)


# CHECKING THAT THE FUNCTION IS WORKING AS IT SHOULD
# check that the dummy train corpus is segmented the way it should:

assert_equal(segment_text(dummy_train_path, dummy_merges), ['low',
                                                             'low',
                                                             'low',
                                                             'low',
                                                             'low',
                                                             'low e s t',
                                                             'low e s t',
                                                             'new er',
                                                             'new er',
                                                             'new er',
                                                             'new er',
                                                             'new er',
                                                             'new er',
                                                             'W I D er',
                                                             'W I D er',
                                                             'W I D er',
                                                             'new',
                                                             'new'])


# check that the dummy test corpus is segmented the way it should:
dummy_test_path = "/coursedata/06-subwords/dummy_test_corpus.txt"
assert_equal(segment_text(dummy_test_path, dummy_merges), ['low er', 'c o o l er'])


```

## TASK 2 <a class="anchor" id="task_2"></a>
## ANALYZE SEGMENTATIONS
## 2.1 <a class="anchor" id="subtask_2_1"></a>
### Count word OOV (1 Point)
Now that we've done with the algorithm, let's see if it will actually help us with the OOV problem.
Let's use the same corpus as we did in the previous POS-tagging assignment. We've randomly shuffled the sentences and split the corpus in roughly half. One half will be our training example, and another will be our test example.

Analyze:
1. The number of word **types** in the vocabulary of the training corpus.
2. The number of word **types** in the vocabulary of the test corpus.
3. The percentage of word **types** in the test corpus vocabulary that are absent in the training corpus. (what part of test corpus vocabulary is OOV?)

Run the cell below, to collect the tokenized corpora (we're splitting the words in the corpora by whitespaces), type in the answer in the next cell.

```python
with open("/coursedata/06-subwords/gum_train.txt", 'r') as f:
    train = f.read()
    train = train.split()

with open("/coursedata/06-subwords/gum_test.txt", 'r') as f:
    test = f.read()
    test = test.split()
```

```python
set1=set(train)
set2=set(test)
print(len(set1))
print(len(set2))
print((len(set2)-len(set1.intersection(set2)))/(len(set1))*100)
```

Output:
```
8677
8770
58.026967846029734
```

<!-- grade_id: cell-6b5793a004cb3f7c, points: 0, grade: false, solution: true -->
```python
# type in the answer as integer number
# For example:
# len_of_train_vocab = 234
len_of_train_vocab = 8677

# type in the answer as integer number
# For example:
# len_of_test_vocab = 234
len_of_test_vocab = 8770

# type in the answer as a float number between 0 and 100
# For example:
# oov_percentage = 90.9
oov_percentage= 58.03

#Remember to remove the raise NotImplementedError line:
# YOUR CODE HERE
```

<!-- grade_id: cell-dc2b01fa475e484e, points: 0.3, grade: true, solution: false -->
```python
from numpy.testing import assert_almost_equal
from nose.tools import assert_equal

# QUESTION 1 tests

# checks if your answer is an integer
assert_equal(type(len_of_train_vocab), int)
# checks if your answer is remorely correct integer
assert(7000 < len_of_train_vocab < 10000)

```

<!-- grade_id: cell-f522e06e0c3478c1, points: 0.3, grade: true, solution: false -->
```python
# QUESTION 2 tests

# checks if your answer is an integer
assert_equal(type(len_of_test_vocab), int)
# checks if your answer is remorely correct integer
assert(7000 < len_of_test_vocab < 10000)

```

<!-- grade_id: cell-1223f6fd86f2005c, points: 0.4, grade: true, solution: false -->
```python
# QUESTION 2 tests

# checks if your answer is a float
assert_equal(type(oov_percentage), float)
# checks if your answer is remorely correct integer
assert(1 < oov_percentage < 70)

```

## 2.2 <a class="anchor" id="subtask_2_2"></a>
### Count subword OOV (1 Point)

now Let's compare the OOV numbers we've got in the case where the text is tokenized by words and the case when it's tokenized by subwords. 

Learn 5000 BPE segmentation from the training data, then segment both corpora and compare the vocabulary numbers again. Note that it will take a couple of minutes to run BPE.

Analyze:
1. The number of subword **types** in the vocabulary of the training corpus.
2. The number of subword **types** in the vocabulary of the test corpus.
3. The percentage of subword **types** in the test corpus vocabulary that are absent in the training corpus. (what part of test corpus vocabulary is OOV?)

```python
merges = learn_BPE_merges("/coursedata/06-subwords/gum_train.txt", 5000)
segmented_train = segment_text("/coursedata/06-subwords/gum_train.txt", merges)
segmented_test = segment_text("/coursedata/06-subwords/gum_test.txt", merges)
# represent texts as a list of subwords
segmented_train_as_subwords = " ".join(segmented_train).split(" ")
segmented_test_as_subwords = " ".join(segmented_test).split(" ")


# double checking that your BPE algorithm is working correctly
assert_equal(merges[0], ('e', '_'))
assert_equal(merges[1000], ('t', 'ro'))
assert_equal(merges[-1], ('des', 'cendants_'))

assert_equal(segmented_train[580], 'V ill age')
assert_equal(segmented_train[5400], 'laun ching')

assert_equal(segmented_test[501], 'C a the dr al')
assert_equal(segmented_test[517], 'ho t els')

assert(len(segmented_train_as_subwords) == 64166)
assert(len(segmented_test_as_subwords) == 72712)
```

```python
set1=set(segmented_train_as_subwords)
set2=set(segmented_test_as_subwords)
print(len(set1))
print(len(set2))
print((len(set2)-len(set1.intersection(set2)))/(len(set1))*100)
```

Output:
```
4281
3595
3.1067507591684187
```

<!-- grade_id: cell-e5d749c71ee045a7, points: 0, grade: false, solution: true -->
```python
# type in the answer as an integer number
# For example:
# len_of_train_sub_vocab = 234
len_of_train_sub_vocab = 4281

# type in the answer as an integer number
# For example:
# len_of_test_sub_vocab = 234
len_of_test_sub_vocab = 3595

# type in the answer as a float number between 0 and 100
# For example:
# oov_sub_percentage = 50.9
oov_sub_percentage  = 3.11

#Remember to remove the raise NotImplementedError line:
# YOUR CODE HERE
```

<!-- grade_id: cell-c29d08ffc3d29073, points: 0.3, grade: true, solution: false -->
```python
# QUESTION 1 tests

# checks if your answer is an integer
assert_equal(type(len_of_train_sub_vocab), int)
# checks if your answer is remorely correct integer
assert(500 < len_of_train_sub_vocab < 5000)

```

<!-- grade_id: cell-786b076f34bd1f3f, points: 0.3, grade: true, solution: false -->
```python
# QUESTION 2 tests

# checks if your answer is an integer
assert_equal(type(len_of_test_sub_vocab), int)
# checks if your answer is remorely correct integer
assert(500 < len_of_test_sub_vocab < 5000)

```

<!-- grade_id: cell-3d0f1cd7e311c960, points: 0.4, grade: true, solution: false -->
```python
# QUESTION 3 tests

# checks if your answer is a float
assert_equal(type(oov_sub_percentage), float)
# checks if your answer is remorely correct integer
assert(1 < oov_sub_percentage < 50)

```

## 2.3 <a class="anchor" id="subtask_2_3"></a>
### Does the segmentation make sense? (1 Point)
Now let's look a bit closer at the subwords that we've learned. Take a second to think if the results are what you would expect them to be.

1. What are the top 10 most frequent subwords in the segmented test corpus? (in decending frequency order)
2. What are the top 5 longest subwords in the segmented test corpus? (sort alphabetically)
3. What are the top 5 most frequent lengths of subwords in the test corpus?

<!-- grade_id: cell-e30eb3fbd26f4157, points: 0, grade: false, solution: true -->
```python
# type in the answer as a list of strings
# For example:
# top_10_by_freq = ['a','b'...]
top_10_by_freq = ['the',',','.','of','and','a','in','to','on','is']

# type in the answer as a list of strings
# For example:
# top_5_by_len = ['a','b'...]
top_5_by_len = ['entertainment', 'international', 'unprecedented', 'representative', 'transportation']

# type in the answer as a list of integers
# For example:
# top_5_freqs_of_lens= [1,2,3...]
top_5_freqs_of_lens = [2, 1, 3, 4, 5]

#Remember to remove the raise NotImplementedError line:
# YOUR CODE HERE
```

```python
counts=Counter(segmented_test_as_subwords)
print(counts.most_common(10))
print(sorted(sorted(set(segmented_test_as_subwords)),key=len)[-10:])
counts_3=Counter(list(map(len,(segmented_test_as_subwords))))
print(counts_3.most_common(5))
```

Output:
```
[('the', 2456), (',', 2295), ('.', 2249), ('of', 1356), ('and', 1301), ('a', 1248), ('in', 1138), ('to', 1075), ('on', 527), ('is', 522)]
['specifically', 'International', 'Revolutionary', 'approximately', 'concentration', 'entertainment', 'international', 'unprecedented', 'representative', 'transportation']
[(2, 19569), (1, 17039), (3, 16870), (4, 9100), (5, 4412)]
```

<!-- grade_id: cell-e7786799f7aeacdc, points: 0.3, grade: true, solution: false -->
```python
# QUESTION 1 tests

# checks if your answer is a list of strings
assert_equal(type(top_10_by_freq), list)
assert_equal(type(top_10_by_freq[0]), str)
assert_equal(len(top_10_by_freq), 10)

```

<!-- grade_id: cell-7f4032e3c0f62c07, points: 0.3, grade: true, solution: false -->
```python
# QUESTION 2 tests

# checks if your answer is a list of strings
assert_equal(type(top_5_by_len), list)
assert_equal(type(top_5_by_len[0]), str)
assert_equal(len(top_5_by_len), 5)

```

<!-- grade_id: cell-68ac266fbfcf1cef, points: 0.4, grade: true, solution: false -->
```python
# QUESTION 3 tests

# checks if your answer is a list of integers
assert_equal(type(top_5_freqs_of_lens), list)
assert_equal(type(top_5_freqs_of_lens[0]), int)
assert_equal(len(top_5_freqs_of_lens), 5)


```

## 2.4 <a class="anchor" id="subtask_2_4"></a>
### Your thoughts (3 Points)
Briefly answer the following questions:

1. Describe what will happen if you change the k parameter? How to find a good number for k?

2. What are the possible NLP applications that can benefit from the tokenization into subwords? Why?

3. How the OOV number for our data can be lowered further without changing anything in the segmentation procedure (k stays the same)?

<!-- grade_id: cell-18c675e80eea69ed, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

1. If k is too small, then the subwords will be too short and cannot get long and accurate subwords. If k is too large, then the computation is heavy. To choose a good k, I would say a possible way is to use the percentage of subword types in the test corpus vocabulary that are absent in the training corpus as a measure.

2. An example use case is neural machine translation (NMT). It makes the NMT model capable of open-vocabulary translation by encoding rare and unknown words as sequences of subword units. As we can see in this exercise, the out-of-vocabulary is dramatically decreased using BPE. So BPE helps the translation of out-of-vocabulary words.

3. Enlarge training data set to cover more words.

## Checklist before submission <a class="anchor" id="checklist"></a>
### 1
To make sure that you didn't forget to import some package or to name some variable, press **Kernel -> Restart** and then **Cell -> Run All**. This way your code will be run exactly in the same order as during the autograding.
### 2
Click the **Validate** button in the upper menu to check that you haven't missed anything. You might need to run the validation in the terminal because BPE algorithm takes time.
### 3
To submit the notebook, click on the **jupyterhub** logo in the upper left part of the window, choose the **Assignments** folder, and press **submit**. You can submit multiple times, only the last one counts.
### 4
Please fill in the feedback form in the [Assignment](https://mycourses.aalto.fi/mod/questionnaire/view.php?id=718490) section of Mycoures.